<h1 class="alert alert-success">Les dictionnaires</h1>

<h2 class="alert alert-info">Rappels sur les dictionnaires</h2>

Dans cette partie, nous allons fabriquer un carnet d'adresse pour stocker des contacts. 

<h3 class="alert alert-warning">Fabrication d'un contact </h3>

Chaque contact sera un dictionnaire dont les clés seront :
- `nom` : Nom et prénom du contact
- `tel` : N° de téléphone
- `rue` : adresse complète
- `code` : code postal
- `ville` : ville
- `naissance` : date de naissance

Créez un dictionnaire nommé `contact` correspondant au contact suivant :
> Costa Brigitte<br/>
> 08 06 18 37 28<br/>
> 93, avenue Bruneau<br/>
> 13749 Perrot

<!-- grade_id: cell-43440f0b40a6d000, points: 0, grade: false, solution: true -->
```python
# à vous de jouer
# ...

```

<!-- grade_id: cell-8724965a83c0e1e6, points: 1, grade: true, solution: false -->
```python
# Vérification
assert contact["nom"] == "Costa Brigitte"
assert contact["tel"] == "08 06 18 37 28"
assert contact["ville"] == "Perrot"
```

Ajouter une nouvelle entrée "naissance" dans le contact ayant pour valeur '1987-02-23'

<!-- grade_id: cell-d25894b093958cf5, points: 0, grade: false, solution: true -->
```python
# à vous de jouer
# ...

```

<!-- grade_id: cell-914fdb244f34c260, points: 1, grade: true, solution: false -->
```python
# Vérification
assert contact["naissance"] == '1987-02-23'
```

<h3 class="alert alert-warning">Génération automatique d'un carnet de contacts</h3>

On dispose d'un fichier binaire annexe contenant une liste-Python de 2048 contacts : `P1-3-carnet-2048`.

Chaque contact est un dictionnaire dont les clefs sont : nom, tel, rue, code, ville, naissance. 

La lecture de ce fichier et l'extraction de données peut se faire à l'aide du module `pickle` comme le montre l'exemple ci-dessous.

```python
import pickle
```

```python
with open('P1-3-carnet-2048', 'rb') as fichier:  # ouverture du fichier en lecture 'rb' (read binary)
    carnet = pickle.load(fichier)                # extraction de l'objet liste-Python contenue dans le fichier

# affichage des 3 premiers contacts de la liste :
carnet[:3]
```

#### Remarque : le fichier de contacts a été créé initialement à l'aide du module `faker`.
On donne ci-dessous un exemple d'utilisation de ce module (attention : non disponible dans Basthon).

```python
from faker import Faker

fake = Faker("fr_FR") # Générateur de données personnelles pour un Français

print(fake.last_name())
print(fake.first_name())
print(fake.phone_number())
print(fake.street_address())
print(fake.postcode(), fake.city())
print(fake.date())
```

Tout est à présent en place pour que nous puissions fabriquer notre carnet d'adresse.

**Première implémentation**

Dans une première approche, nous allons considérer que le carnet d'adresse sera une liste de contacts, chaque contact étant un dictionnaire dont la structure a été définie précédemment.

**Écrire une fonction `genere_carnet1`**
- prenant en paramètre le nombre `n` de contacts à générer.
- renvoyant une **liste** de `n` contacts générés aléatoirement.

*Remarque* : on pourra utiliser la fonction `shuffle` du module `random` qui permet de mélanger une liste-Python.

```python
from random import shuffle
help(shuffle)
```

<!-- grade_id: cell-e6831bf13c8a8f2a, points: 0, grade: false, solution: true -->
```python
def genere_carnet1(n):
    """Renvoie une liste de n contacts aléatoires"""
    assert n <= 2048, "taille maxi du carnet : 2048 contacts."
    # à vous de jouer
    # ...

```

```python
carnet = genere_carnet1(5)
carnet
```

**Révision au passage... : tri d'un tableau en Python**

```python
for contact in sorted(carnet, key=lambda dico:dico['naissance']):
    print(contact['naissance'], contact['nom'])
    
print()

for contact in sorted(carnet, key=lambda dico:dico['nom']):
    print(contact['nom'], contact['naissance'])
```

<!-- grade_id: cell-1c391af7665fa569, points: 0, grade: false, solution: true -->
```python
def est_present(nom, carnet):
    """Teste si nom est présent dans le carnet d'adresse"""
    # à vous de jouer
    # ...

```

<!-- grade_id: cell-4e957defe3520970, points: 1, grade: true, solution: false -->
```python
# Vérification
carnet1 = genere_carnet1(10)
nom = carnet1[0]["nom"]
assert est_present(nom, carnet1)
assert not est_present("Longtemps Adam", carnet1)
```

<h3 class="alert alert-warning">Mesure de performance de la recherche</h3>

Nous allons regarder ici comment évolue la vitesse de recherche en fonciton de la taille du carnet d'adresse. On utilisera pour ncela la fonction magique de *jupyter* : `%%timeit`.
Etudiez la cellule suivante :

```python
# Fabrication d'un carnet de 100 contacts
carnet1 = genere_carnet1(100)
nom = carnet1[-1]["nom"]  # On récupère un nom du carnet
nom
```

```python
%%timeit

# On mesure le temps d'une recherche
est_present(nom, carnet1) 
```

Vous lisez sous la cellule le temps de recherche.

À présent, on refait l'expérience pour 1000 contacts dans le carnet d'adresse.

```python
carnet1 = genere_carnet1(1000)
nom = carnet1[-1]["nom"]  # On récupère un nom du carnet
nom
```

```python
%%timeit

# On mesure le temps d'une recherche dans ce carnet
est_present(nom, carnet1)
```

<h3 class="alert alert-warning">Seconde implémentation</h3>

Vous devez avoir constaté ci-dessus que le temps de recherche est proportionnel à la taille du carnet d'adresse : si celui-ci contient 10 fois plus de contact, la recherche peut être jusqu'à 10 fois plus longue.

Nous allons changer d'approche et fabriquer un carnet d'adresse sous forme d'un dictionnaire dont les clés seront les **noms** et les valeurs seront les fiches contacts. Ainsi notre carnet d'adresse sera un dictionnaire dont les valeurs seront des dictionnaires !

**Écrire une fonction `genere_carnet2`**
- prenant en paramètre le nombre n de contacts à générer
- renvoyant un **dictionnaire** de `n` contacts générés aléatoirement : les clefs sont les noms des contacts.

*Remarque :* On pourra d'abord générer une liste (avec la fonction `genere_contact1`) puis chaque contact de la liste devient une **valeur** du dictionnaire dont la **clef** est simplement le nom du contact.

<!-- grade_id: cell-94349ffa9c7ab8c0, points: 0, grade: false, solution: true -->
```python
def genere_carnet2(n):
    """Renvoie un dictionnaire de n contacts aléatoires"""
    # à vous de jouer
    # ...

```

```python
carnet = genere_carnet2(5)
carnet
```

<h3 class="alert alert-warning">Mesure de performance de la recherche</h3>

Nous allons regarder pour cette nouvelle implémentation comment évolue la vitesse de recherche en fonction de la taille du carnet d'adresse. Validez les 2 cellules suivantes.

```python
# Fabrication d'un carnet de 100 contacts
carnet = genere_carnet2(100)
nom = list(carnet.keys())[-1] # On récupère un nom du carnet
nom
```

```python
%%timeit

nom in carnet                 # On le recherche
```

On constate déjà que la recherche est plus rapide que pour la première implémentation du carnet à l'aide d'un tableau.

Refaisons l'expérience avec 10 fois plus de contacts dans le carnet !!

```python
# Fabrication d'un carnet de 1000 contacts
carnet = genere_carnet2(1000)
nom = list(carnet.keys())[-1] # On récupère un nom du carnet
nom
```

```python
%%timeit

nom in carnet           # On le recherche
```

<h2 class="alert alert-warning">Conclusion</h2>
<div class="alert alert-danger">
Vous le constatez d'après les expériences ci-dessus : le temps de recherche dans le dictionnaire est pratiquement indépendant du nombre d'entrées dans ce dictionnaire, car en multipliant le nombre de contacts par 100, le temps est resté pratiquement identique alors que dans le cas de la recherche dans un tableau, celui-ci est proportionnel à la longueur du tableau.
<br><br>
Le dictionnaire est donc une structure de données optimisée pour la recherche sur les clés.
    </div>

<h3 class="alert alert-warning">Pour aller plus loin</h3>
On peut recommencer ces calculs de temps d'exécution pour des collections de contacts de tailles diverses et tracer le graphique *duree=f(taille)*.

1. Mesurer le temps d'exécution de chaque recherche pour des collections de 1 à 1024 contacts, en doublant à chaque fois la taille de la collection de contacts.
2. Relever toutes les durées et tracer le graphique duree=f(taille).
- Pour cela, importer le module matplotlib.pyplot (sous l'alias *plt*)
- la fonction pour tracer un graphique est ```plot(x, y)``` : x et y étant une liste d'abscisses et une liste d'ordonnées.

```python
from time import time
```

```python
# à vous de jouer
tailles = [2**i for i in range(10)]               # Définition des tailles
carnets1 = [genere_carnet1(n) for n in tailles]   # création des carnets sous formes de listes
carnets2 = [genere_carnet2(n) for n in tailles]   # créations des carnets sous forme de dictionnaire

duree_liste, duree_dico = [], []                  # pour enregistrer les durées de recherche en fonction des tailles

nom = ''                                          # nom qui est forcément absent de la collection de contacts

for carnet in carnets1:
    debut = time()                                # enregistre la 'date' de début des calculs
    for i in range(100):                          # on fait 100 fois le calcul pour moyenner
        test = est_present(nom, carnet)
    duree_liste.append(time() - debut)            # ajout de la durée de calcul à la liste
    
for carnet in carnets2:
    # à vous de jouer pour enregistrer les durées de recherche dans les dictionnaires
    # ...
```

```python
import matplotlib.pyplot as plt
```

```python
# TRACER LES GRAPHIQUES
# à vous de jouer
# ...

```

```python

```