# Algorithmes de tri

<div class="alert alert-success">

# Tris par insertion et par sélection, complexité
</div>

<h2 class="alert alert-info">Tri par sélection</h2>

<div class="alert alert-warning">
    
Compléter la fonction `tri_selection(t)` qui trie en place un tableau `t` (ne renvoie rien) qui implémente l'algorithme de tri par sélection.
    
</div>

<!-- grade_id: cell-fbb2ff38ded373b2, points: 0, grade: false, solution: true -->
```python
def tri_selection(t):
    """ Tri un tableau t d'entiers.
    Entrée : t : list(int)
    Sortie : None (remarque : c'est un tri en place)
    """
    for i in range(...):
        i_mini = i
        for j in range(..., ...):
            if ... < ...:
                i_mini = j
        t[i], t[i_mini] = t[...], t[...]
```

<!-- grade_id: cell-fbb2ff38ded373b2, points: 0, grade: false, solution: true -->
```python
def tri_selection(t):
    """ Tri un tableau t d'entiers.
    Entrée : t : list(int)
    Sortie : None (remarque : c'est un tri en place)
    """
    for i in range(len(t)-1):
        i_mini = i
        for j in range(i+1, len(t)):
            if t[j] < t[i_mini]:
                i_mini = j
        t[i], t[i_mini] = t[i_mini], t[i]
```

<div class="alert alert-warning">
    
**Test :** 
    
1. Créer un tableau `tab` de 20 entiers compris entre 1 et 30.  (utlisation de la fonction `randint` du module random)
2. Trier le tableau avec la fonction de tri par sélection et vérifier que le tableau est bien trié.
    
</div>

```python
from random import randint 
# randint(a, b) renvoie un nombnre entier aléatoire entre a et b (inclus)
```

<!-- grade_id: cell-7020125391792ebf, points: 0, grade: false, solution: true -->
```python
# A VOUS DE JOUER
tab = [randint(1, 30) for i in range(20)]
tab
```

```python
tri_selection(tab)
tab
```

<h2 class="alert alert-info">Tri par insertion</h2>

<div class="alert alert-warning">
    
Compléter la fonction `tri_insertion(t)` qui trie en place un tableau `t` (ne renvoie rien) qui implémente l'algorithme de tri par insertion.
    
</div>

<!-- grade_id: cell-11c9879c87845f89, points: 0, grade: false, solution: true -->
```python
def tri_insertion(t):
    """ Tri un tableau t d'entiers.
    Entrée : t : list(int)
    Sortie : None (remarque : c'est un tri en place)
    """
    for i in range(..., ...):
        element = t[i]
        j = i - 1
        while (... >= ...) and (t[j] ... element):
            t[j+1] = t[...]
            j = ...
        t[...] = element
```

<!-- grade_id: cell-11c9879c87845f89, points: 0, grade: false, solution: true -->
```python
def tri_insertion(t):
    """ Tri un tableau t d'entiers.
    Entrée : t : list(int)
    Sortie : None (remarque : c'est un tri en place)
    """
    for i in range(1, len(t)):
        element = t[i]
        j = i - 1
        while (j >= 0) and (t[j] > element):
            t[j+1] = t[j]
            j = j - 1
        t[j+1] = element
```

<div class="alert alert-warning">
    
**Test :** 
    
1. Créer un tableau `tab` de 20 entiers aléatoires compris entre 1 et 30.
2. Trier le tableau avec la fonction de tri par sélection et vérifier que le tableau est bien trié.
    
</div>

<!-- grade_id: cell-7020125391792ebf, points: 0, grade: false, solution: true -->
```python
# A VOUS DE JOUER
tab = [randint(1, 30) for i in range(20)]
tab
```

```python
tri_insertion(tab)
tab
```

<h2 class="alert alert-info">Complexité : vérification du coût quadratique</h2>

<div class="alert alert-warning">
    
Créer un tableau `tab` de 20 entiers compris entre 1 et 30.
    
</div>

<!-- grade_id: cell-7020125391792ebf, points: 0, grade: false, solution: true -->
```python
# A VOUS DE JOUER
t = [randint(1, 30) for i in range(20)]
```

<!-- grade_id: cell-f8dfbe7b8daf9252, points: 2, grade: true, solution: false -->
```python
tri_insertion(t)
t
```

```python
t = [randint(1, 30) for i in range(20)]
print(t)
```

<!-- grade_id: cell-eedafefd27515f77, points: 2, grade: true, solution: false -->
```python
tri_python = sorted(t)
assert tri_python == t
```

On désire vérifier que le coût de ces algorithmes est quadratique. Pour cela, nous allons exécuter les algorithmes sur des tableaux de tailles variables et mesurer le temps d'exécution dans chaque cas.

<div class="alert alert-warning">
    
Lire l'instruction suivante et expliquer ce que contient le tableau `data` créé.
    
</div>

<!-- grade_id: cell-d6416617e10d311c, points: 0, grade: false, solution: true -->
```python
data = [[randint(0, 1000) for j in range(1000 * i)] for i in range(1, 11)]
```

*Votre réponse ici :*

...

...

Le module `time` fournit une fonction `time()` qui renvoie le temps  écoulé depuis le 1er janvier 1970, en secondes.

Documentation : https://docs.python.org/fr/3/library/time.html

On peut donc mesurer le temps d'exécution d'une série d'instructions de la manière suivante :

    debut = time()
    < suite d'instructions >
    fin = time()
    duree = fin - debut
    

```python
from time import time
```

Relever dans deux listes `durees_insertion` et `durees_selection` les durée d'exécution des algorithmes de tri pour chacun des tableaux de `data`.

**ATTENTION** : il faut ***re-générer la liste `data`*** après la première série de tris, sinon tous les tableaux sont déjà triés (puisque le tri se fait en place) avant d'attaquer la 2ème série de tris !

<!-- grade_id: cell-903eb3a2feaa3a47, points: 3, grade: true, solution: true -->
```python
# le travail est djà écrit pour le tri par insertion :

data = [[randint(0, 1000) for j in range(1000*i)] for i in range(1, 11)]

durees_insertion = []

for i in range(10):
    debut = time()
    tri_insertion(data[i])
    fin = time()
    durees_insertion.append(fin-debut)    
```

<div class="alert alert-warning">
    
Réprendre le même type de calcul de durées pour le tri par sélection.
    
</div>

<!-- grade_id: cell-903eb3a2feaa3a47, points: 3, grade: true, solution: true -->
```python
data = [[randint(0, 1000) for j in range(1000*i)] for i in range(1, 11)]

durees_selection = []

# A VOUS DE JOUER
# 
# ...
```

<!-- grade_id: cell-903eb3a2feaa3a47, points: 3, grade: true, solution: true -->
```python
data = [[randint(0, 1000) for j in range(1000*i)] for i in range(1, 11)]

durees_selection = []

# A VOUS DE JOUER
# Réprendre le même type de calcul de durée pour le tri par sélection

for i in range(10):
    debut = time()
    tri_selection(data[i])
    fin = time()
    durees_selection.append(fin-debut)    
```

```python
# Affichage des durées calculées
for i in range(10):
    print(f"tableau de taille n = {(1+i)*1000}")
    print(f"insertion : {durees_insertion[i]}")
    print(f"selection : {durees_selection[i]}\n")
```

<h2 class="alert alert-info">Complexité : Représentation graphique</h2>

Nous allons maintenant représenter des graphiques donnant la durée d'exécution des algorithme en fonction de la taille des tableaux à trier.

Puis nous comparerons l'allure de ces graphique à la fonction carré ($f(x) = x^2$).

Pour cela, nous aurons besoin du module `matplotlib.pyplot` pour tracer un graphique.

```python
import matplotlib.pyplot as plt
```

Observer bien l'exemple suivant pour découvrir les fonctionnalités de base (cela peut-être beaucoup plus compliqué !) du module `matplotlib.pyplot`.

```python
# création des données :
abscisses = [i for i in range(1, 11)]
ordonnees = [x**2 for x in abscisses]

# création du graphique
plt.plot(abscisses, ordonnees, color='green', label='fonction carré') # prépare le graphique
plt.legend()  # commande l'affichage de la légende (càd le label défini précédemment)
plt.show()    # présente le graphique à l'écran
```

<div class="alert alert-warning">
    
En vous appuyant sur le modèle de la cellule précédente, tracez les graphiques donnant la durée d'exécution des algorithmes en fonction de la taille des données.
</div>

<!-- grade_id: cell-88d4ab71c04bfc92, points: 3, grade: true, solution: true -->
```python
# création des abscisses :
n = [n for n in range(1, 11)]

# A VOUS DE JOUER

# création du graphique sélection
plt.plot(n, durees_selection, color='red', label='tri selection') 
plt.legend()
plt.show()

# création du graphique insertion
plt.plot(n, durees_insertion, color='blue', label='tri insertion') 
plt.legend() 
plt.show() 
```

Conclusion : que pensez-vous des allures des graphiques obtenus ? Le coût semble-t-il quadratique ?

*Votre réponse ici :*

...

...

Ci-dessous, on présente les courbes sur un unique graphique pour comparer plus aisément les courbes entre elles. 

Par ailleurs, on a normé les résultats pour éviter des soucis d'échelle.

```python
n = [n for n in range(1, 11)]

durees_normees_selection = [duree/durees_selection[-1] for duree in durees_selection]
durees_normees_insertion = [duree/durees_insertion[-1] for duree in durees_insertion]
carres_normes = [i**2/100 for i in range(1, 11)]

plt.plot(n, durees_normees_selection, color='red', label='selection')
plt.plot(n, durees_normees_insertion, color='blue', label='insertion')
plt.plot(n, carres_normes, color='black', label='$f(x)=x^2$')
plt.legend()
plt.show()
```

```python

```